Setting Up a ReadOnly Database Connection in Pandas
-----------------------------------------------------------------

This assignment will make sure that you have the database setup by testing to see if 
we can make a database connection to the shared database for the class.  It also will test
if your database connection information (sql_string) is set properly for the autograder.

```python
# https://www.pg4e.com/code/ipynb/01-readonly-python.ipynb

import psycopg2
import pandas as pd
```

```python
# If you are going to send this file to the autograder, you cannot use hidden.py
# and must uncomment the sql_string assignment statement below and put in your values

sql_string = None
# sql_string = 'dbname=pg4e_data user=pg4e_data_read password=pg4e_pass_94e5d host=35.239.113.162 port=10001'
    
# If we leave sql_string as None, we can use hidden.py - but this only works for demo/test - not autograder
if sql_string is None:
    import hidden
    secrets = hidden.readonly()
    sql_string = hidden.psycopg2(hidden.readonly())
    print('PostgreSQL connection data taken from hidden.py')
```

Output:
```
PostgreSQL connection data taken from hidden.py
```

```python
conn = psycopg2.connect(sql_string,connect_timeout=3)
```

```python
sql = "SELECT count(*) AS count FROM taxdata;"
df = pd.read_sql_query(sql, conn)
assert df['count'][0] == 378304
```

```python
df.head()
```

Output:
```
count
0  378304
```

<!-- grade_id: cell-ad191fd62ca88a09, points: 1, grade: true, solution: false -->
```python
my_sql = 'select * from taxdata where ein = 742661023 limit 1;'
my_df = pd.read_sql_query(my_sql, conn)
assert my_df['expenses'][0] == 0
assert my_df['revenue'][0] == 0

### BEGIN HIDDEN TESTS
my_sql = 'select * from taxdata where ein = 270243010 limit 1;'
my_df = pd.read_sql_query(my_sql, conn)
assert my_df['expenses'][0] == 147332
assert my_df['revenue'][0] == 153662
### END HIDDEN TESTS
```

```python

```