Reading Normalized Data Quickly using a Database
------------------------------------------------------------

In your readonly database, there is an un-normalized table called *home_value_by_zip* with 4,466,776 records of un-normalized data.  There is also a normalized copy of the same data in a few tables in the database.  This goal of this assignment is to use JOIN and GROUP BY to query the normalized data, resulting in faster and more efficient queries.

You must examine the tables using psql and figure out how to connect them and then construct an efficient SQL query in this notebook that will retrieve the requested data.

```python
# https://www.pg4e.com/code/ipynb/03-normalization.ipynb

import psycopg2
import pandas as pd
import time
```

```python
# Update this SQL connection string with correct readonly database credentials
sql_string = 'dbname=readonly user=readonly password=PLEASE_UPDATE_ME host=pg.pg4e.com port=5432'
```

```python
# open a new connection
# NOTE: this will raise an OperationalError if the readonly database credentials are not updated in the previous cell
conn = psycopg2.connect(sql_string,connect_timeout=3)
print(f"connection created: {conn}")
```

You are to construct a query using the normalized tables that will return the same results as:

    SELECT state, avg(ym_val) AS average FROM home_value_by_zip
    GROUP BY state ORDER BY average DESC LIMIT 10;

This query may run quite long - and might not finish - you should construct a more efficient query using JOIN starting with the `home_value` table.

Here is the expected output of the first few rows of the query:

     state |       average       
    -------+---------------------
     CA    | 429388.882710557533
     HI    | 384304.615036999379
     DC    | 373415.607524148449
     NJ    | 313458.077439427195


<!-- grade_id: cell-715a65927576adb6, points: 0, grade: false, solution: true -->
```python
state_avg_sql = None

### DEFINE YOUR SQL QUERY HERE (NOTE: this query will fail as-is)
state_avg_sql = '''
SELECT
    state,
    avg(ym_val) AS average
FROM home_value 
-- add JOIN and GROUP BY logic here...
ORDER BY average
DESC LIMIT 10;
'''
### END SOLUTION

if state_avg_sql is None:
    raise Exception('Please define the sql query above')

start_time = time.time()
state_avg_df = pd.read_sql_query(state_avg_sql, conn)
print(f"state average execution time: {round(time.time()-start_time, 2)}s")
state_avg_df.head()
```

<!-- grade_id: cell-e94015b31a39e6d0, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert state_avg_df['state'][1] == 'HI'
assert state_avg_df['average'][1] > 384304
assert state_avg_df['average'][1] < 384305
assert state_avg_df['state'][3] == 'NJ'
assert state_avg_df['average'][3] > 313458
assert state_avg_df['average'][3] < 313459

if state_avg_sql.lower().find('join') < 0 :
    raise Exception('You need to have a JOIN in your query')
### END HIDDEN TESTS
```

You are to construct a query using the normalized tables that will return the same results as:

    SELECT city, avg(ym_val) AS average FROM home_value_by_zip
    GROUP BY city ORDER BY average DESC LIMIT 10;

This query may run quite long - and might not finish - you should construct a more efficient query using JOIN starting with the `home_value` table.

Here is the expected output of the first few rows of the query:

           city       |       average        
     -----------------+----------------------
      Atherton        | 3625292.526690391459
      Portola Valley  | 2218466.548042704626
      Fisher Island   | 2078791.814946619217
      Montecito       | 1939405.693950177936


<!-- grade_id: cell-fb2f81aa54d0b8cd, points: 0, grade: false, solution: true -->
```python
city_avg_sql = None

### DEFINE YOUR SQL QUERY HERE (NOTE: this query will fail as-is)
city_avg_sql = '''
SELECT
    city,
    avg(ym_val) AS average
FROM home_value 
-- add JOIN and GROUP BY logic here...
ORDER BY average
DESC LIMIT 10;
'''
### END SOLUTION

if city_avg_sql is None:
    raise Exception('Please define the sql query above')

start_time = time.time()
city_avg_df = pd.read_sql_query(city_avg_sql, conn)
print(f"state average execution time: {round(time.time()-start_time, 2)}s")
city_avg_df.head()

```

<!-- grade_id: cell-dbf5ff3e57889dc7, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert city_avg_df['city'][1] == 'Portola Valley'
assert city_avg_df['average'][1] > 2218466
assert city_avg_df['average'][1] < 2218467
assert city_avg_df['city'][3] == 'Montecito'
assert city_avg_df['average'][3] > 1939405
assert city_avg_df['average'][3] < 1939406

if city_avg_sql.lower().find('join') < 0 :
    raise Exception('You need to have a JOIN in your query')
### END HIDDEN TESTS
```